#### Bi 410/510 (Fall 2019)

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

**Group Members**

If this is a group submission edit this cell to add the names and e-mail addresses of the othrer people who worked on the project.
* Peregrine Painter 951521340
* Alex Wong 951549455
* Dorian Yeh 951465043

```python

```

# <span style="color:seagreen;">Project 4: &nbsp; Lists and Loops</span>

###  <span style="color:seagreen">Instructions</span>

The exercises in this notebook require you to write the definitions of a set of Python functions. Each exercise has several notebook cells:
* the function you need to write is described in a set of markdown cells
* next is a code cell where you should enter the Python code for your function definition
* following that are one or more "test" cells, code cells that contain calls to your function

As you write the code for your functions you can execute the test cells to see if your code is working.  Feel free to add additional tests if you want.

Submit the final notebook via Canvas.  **If you worked in a group only one person has to submit the project (but make sure everyone's name is in the first cell).**

###  <span style="color:seagreen">Exercise 4.1: </span>`mean_length`

Implement a function named `mean_length` that will compute the average (mean) length of set of strings.  The argument passed to the function is a list of string objects, and the return value should be the mean length of the strings.

**Note:** If the list is empty the function should return 0.

<!-- grade_id: mean_length, points: 10, grade: true, solution: true -->
```python
def mean_length(lst):
    total_l = 0
    for n in lst:
        total_l += len(n)
#         print(n)
#     print(total_l, itterations)
    if total_l == 0:
        return 0
    else:
        return total_l/len(lst)
#     print(avg_len)
```

##### Sandbox 

Add code cells here to test your function as you develop it.

```python
mean_length(['fee', 'free', 'freed'])
mean_length([])
```

Output:
```
0
```

##### Tests 

<!-- grade_id: mean_length_test1, points: 3, grade: true, solution: false -->
```python
assert mean_length(['fee', 'free', 'freed']) == 4
```

<!-- grade_id: mean_length_test2, points: 3, grade: true, solution: false -->
```python
assert mean_length(['alpha', 'beta', 'gamma', 'delta', 'epsilon']) == 5.2
```

<!-- grade_id: mean_length_test3, points: 4, grade: true, solution: false -->
```python
assert mean_length([]) == 0
```

###  <span style="color:seagreen">Exercise 4.2: </span>`valid_password`

Implement a function named `valid_password` that will be passed a string `s`.  The function return a Boolean value (`True` or `False`) depending on whether `s` is a valid password.  For this project, a valid password:
* must contain at least one letter, one digit, and one punctuation character
* must not contain any other character
* must be at least 8 characters long

This code cell imports some useful strings from Python's `string` library -- you can use them if you want.  See the lecture notes on loops (`Python IV`) for more information.

```python
from string import ascii_letters, digits, punctuation
```

<!-- grade_id: validate, points: 8, grade: true, solution: true -->
```python
def valid_password(s):
    if len(s) < 8:
        return False
    elif not any(n in ascii_letters for n in s) or not any(n in digits for n in s) or not any(n in punctuation for n in s): #1 of
        return False
    elif not all(n in (ascii_letters + digits + punctuation) for n in s): #other symbols
        return False
    else:
        return True

```

##### Sandbox 

Add code cells here to test your function as you develop it.

```python
valid_password("          2!a  ")
print((ascii_letters + digits + punctuation))
s= 'asdaef34124124!!!😒!'
all(n in (ascii_letters + digits + punctuation) for n in s)
```

Output:
```
abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ0123456789!"#$%&'()*+,-./:;<=>?@[\]^_`{|}~
False
```

##### Tests 

<!-- grade_id: vp_1, points: 2, grade: true, solution: false -->
```python
assert valid_password('Password0.0')
```

<!-- grade_id: vp_2, points: 2, grade: true, solution: false -->
```python
assert not valid_password('JSC')         # too short
```

<!-- grade_id: vp_3, points: 2, grade: true, solution: false -->
```python
assert not valid_password('John2019')    # missing punctuation
```

<!-- grade_id: vp_4, points: 2, grade: true, solution: false -->
```python
assert not valid_password('John!?xx')    # missing digit
```

<!-- grade_id: vp_5, points: 2, grade: true, solution: false -->
```python
assert not valid_password('2019#$%^')    # missing letter
```

<!-- grade_id: vp_6, points: 2, grade: true, solution: false -->
```python
assert not valid_password('DumbRul3!😒')    # can't have other symbols
```

###  <span style="color:seagreen">Exercise 4.3: </span>`similarity`

Define a function named `similarity` that computes the similarity of two DNA sequences. 
Return `None` if the sequences are not the same length.
If they are the same length, compute the similarity as the percentage of locations that are the same in each sequence. 
If the sequences are identical similarity is 100%, if they differ at every location they are 0% similar.

<!-- grade_id: similarity, points: 12, grade: true, solution: true -->
```python
def similarity(seq1, seq2):
    if len(seq1) != len(seq2): # are sequences equal length?
        return None
    else:
        counter = 0
        for n in range(len(seq1)): # for each letter index:
            if seq1[n] == seq2[n]: # if they are the same,
                counter += 1    # count 1
#                 print(counter)
        return counter/len(seq1)*100 # count/total=%
            
```

##### Sandbox 

Add code cells here to test your function as you develop it.

```python
similarity('ATATATAT','AAAAAAAA')
```

Output:
```
50.0
```

##### Tests 

<!-- grade_id: cell-fd564130bd21c61a, points: 2, grade: true, solution: false -->
```python
assert similarity('ATATATAT','AAAAAAAA') == 50.0
```

<!-- grade_id: cell-e674894c1554ab96, points: 2, grade: true, solution: false -->
```python
assert similarity('AAAAAAA','AAAAAAA') == 100.0
```

<!-- grade_id: cell-9ed1e5589de90016, points: 2, grade: true, solution: false -->
```python
assert similarity('AAAA','CCCC') == 0.0
```

<!-- grade_id: cell-7fb65a90116ee774, points: 2, grade: true, solution: false -->
```python
assert similarity('AAAAAAA', 'CCC') is None
```

###  <span style="color:seagreen">Exercise 4.4: </span>`max_infection`

Suppose researchers have analyzed data from the spread of a particular disease and have fit the data to the following model, which predicts the number of people in a small population who will have the disease on day $t$:

$$
p(t) = \frac{20t^3 - t^4}{1000}
$$

Define a function named `max_infection` that will compute $p(t)$ for 20 days (_i.e._ for $1 \leq t \leq 20$) and return two values:  the day at which the most people are infected, and the number of people infected on that day.

There are no arguments passed to this function.  The expected answer is 16.875 people infected on day 15.

<!-- grade_id: max_infection, points: 12, grade: true, solution: true -->
```python
def max_infection():
    p_max = 0
    for t in range(1,21):
        p_t = (20*t**3-t**4)/1000
        if p_max < p_t:
            p_max = p_t
            day_max = t
            
    return (day_max, p_max)
```

##### Sandbox 

Add code cells here to test your function as you develop it.

```python
max_infection()
```

Output:
```
(15, 16.875)
```

##### Tests 

The expected output is the pair of numbers `(15, 16.875)`:

<!-- grade_id: cell-a5484d2e0169d580, points: 8, grade: true, solution: false -->
```python
assert max_infection() == (15, 16.875)
```

###  <span style="color:seagreen">Exercise 4.5: </span>`codons`

Write a function named `codons` that will return the list of 3-letter substrings of the string passed as an argument.  If the length of the input sequence is not a multiple of 3 the last codon will have less than three letters.

<!-- grade_id: codons, points: 10, grade: true, solution: true -->
```python
from math import ceil
def codons(seq):
    c_list = [] # make empty list
    for n in range(ceil(len(seq)/3)): # itterate for # of codons that can fit, rounding up
        c_list += [seq[n*3:n*3+3]] # make a new list of the iterated codon, and append it to the list of codons
    return c_list
        
```

##### Sandbox 

Add code cells here to test your function as you develop it.

```python
l = l + ["talk"]
print(l)
```

Output:
```
['talk', 'talk', 'talk', 'talk', 'talk', 'talk']
```

```python
codons('GATTACA')
'GATTACA'[6:9]
ceil(len('GATTACA')/3)
```

Output:
```
3
```

##### Tests 

This sequence has 4 codons:  `ATG`, `CTC`, `AAG`, and `TAG`:

<!-- grade_id: cell-4b54c038886eb31c, points: 5, grade: true, solution: false -->
```python
assert codons('ATGCTCAAGTAG') == ['ATG', 'CTC', 'AAG', 'TAG']
```

This sequence has 2 full codons (`GAT` and `TAC`) and one partial codon (`A`):

<!-- grade_id: cell-12ff1d7b61501157, points: 5, grade: true, solution: false -->
```python
assert codons('GATTACA') == ['GAT', 'TAC', 'A']
```

```python

```