#### Bi 410/510 (Fall 2019)

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

**Group Members**

If this is a group submission edit this cell to add the names and e-mail addresses of the othrer people who worked on the project.
* name (xxx@uoregon.edu)
* name (xxx@uoregon.edu)
* name (xxx@uoregon.edu)

# <span style="color:seagreen;">Project 2: &nbsp; Introduction to Python</span>

This is the first Python programming project of the term.  The goal is to learn about basic arithmetic and string operations by defining and testing small function definitions.

###  <span style="color:seagreen">Instructions</span>

The exercises in this notebook require you to write the definitions of a set of Python functions. Each exercise has several notebook cells:
* the function you need to write is described in a set of markdown cells
* next is a code cell where you should enter the Python code for your function definition
* following that are one or more "test" cells, code cells that contain calls to your function

As you write the code for your functions you can execute the test cells to see if your code is working.  Feel free to add additional tests if you want.

Submit the final notebook via Canvas.  If you worked in a group only one person has to submit the project (but make sure everyone's name is in the first cell).

###  <span style="color:seagreen">Exercise 2.0: </span>`vol`

The first project has been completed for you as an example.

The volume of a cylinder with radius $r$ and height $h$ is  

$$
v = h \times \pi \times r^2
$$

Define a function named `vol` that will take two arguments, `h` and `r`, and return the volume of a cylinder with height `h` and radius `r`, rounded to the nearest integer.

```python
from math import pi

def vol(h, r):
    v = h * pi * r**2
    return round(v)
```

##### Tests 

With $h = 8''$ and $r = 1''$ (a tall coffee cup) the volume is around 25 cubic inches (14 ounces).

```python
assert vol(8, 1) == 25
```

A measuring cup $4''$ tall and with a $2\frac{1}{4}''$ radius holds 64 cubic inches (about 4 cups).

```python
assert vol(4, 2.25) == 64
```

###  <span style="color:seagreen">Exercise 2.1: </span>`bmi`

The formula for body mass index, in units used in the U.S., is

$$
\textrm{bmi} = \frac{\textrm{weight}}{\textrm{height}^2} \times 703
$$

Define a function named `bmi` that takes two parameters, a person's `height` and `weight`, and returns their body mass index, rounded to the nearest integer.

The function header has been written for you; you just need to fill in the body of the function.  Delete the line that says `pass` and replace it with the code that computes the BMI.

<!-- grade_id: bmi, points: 6, grade: true, solution: true -->
```python
def bmi(height, weight):
    # YOUR CODE HERE
    b = (weight/height**2)*703
    return round(b)
```

```python
bmi(72, 200)
```

Output:
```
27
```

##### Tests 

A person 6 feet tall weighing 200 pounds has a BMI of 27.

<!-- grade_id: bmi_test1, points: 2, grade: true, solution: false -->
```python
assert bmi(72,200) == 27
```

Someone who is 5' 5" and 140 pounds has a BMI of 23.

<!-- grade_id: bmi_test2, points: 2, grade: true, solution: false -->
```python
assert bmi(65,140) == 23
```

###  <span style="color:seagreen">Exercise 2.2: </span>`pmt`

To determine the monthly payment on a loan first compute a value $r$ using this formula:

$$
r = \mathrm{rate} \, / \, 100 \, / \, 12
$$

where `rate` is the annual interest rate.  Next compute a value $p$ defined as

$$
p = 12 \times \mathrm{years}
$$

where `years` is the length of the loan in years.

Then use the values of $r$ and $p$ in this formula to compute the monthly payment:

$$
\mathrm{payment} = \frac{r \times \mathrm{amt}}{1 - (1 + r)^{-p}}
$$

where `amt` is the loan amount.

Define a Python function named `pmt` that takes three arguments:  a loan amount, an annual interest rate, and the number of years to repay the loan.  The return value should be the monthly payment amount, rounded to two decimal places.

<!-- grade_id: pmt, points: 6, grade: true, solution: true -->
```python
def pmt(amt, rate, years):
    # YOUR CODE HERE
    r = rate/100/12
    p = 12*years
    monthly_payment = (r*amt)/(1-(1+r)**(-p))
    return round(monthly_payment, 2)
```

```python
pmt(12000,15,10)
```

Output:
```
193.6
```

##### Tests 

If a person takes out a 30-year loan for $\$150,000$ at an annual interest rate of $4.5\%$
their monthly payment would be $\$760.03$.

<!-- grade_id: pmt_1, points: 2, grade: true, solution: false -->
```python
assert pmt(150000, 4.5, 30) == 760.03
```

Want to buy a new Tesla Model X?  If you take out a 7-year loan at current interest rates (3.67\%) you would have to pay \$1,196.29 per month.

<!-- grade_id: cell-98d1956b864fc9c4, points: 2, grade: true, solution: false -->
```python
assert pmt(88500, 3.67, 7) == 1196.29
```

###  <span style="color:seagreen">Exercise 2.3: </span>`lcm`

The least common multiple (LCM) of two integers $a$ and $b$ is the smallest number that is evenly divisible by both $a$ and $b$.

For example, suppose $a = 4$ and $b = 6$. The number $48$ is a common multiple because $48$ is evenly divisible by both $4$ and $6$. But the _smallest_ number that is evenly divisible by both $4$ and $6$ is $12$.

It’s simple to compute the LCM of two numbers if we know their greatest common divisor (GCD). The GCD of two numbers is the largest factor they have in common. If we factor the two numbers, $4 = 2 \times 2$ and $6 = 3 \times 2$, and we can see the largest common factor is $2$, so $\mathrm{gcd}(4, 6) = 2$.

Here is the formula for the least common multiple, assuming we know the GCD (https://en.wikipedia.org/wiki/Least_common_multiple).
Assuming $a$ and $b$ are both positive integers:

$$
\mathrm{lcm}(a,b)= \frac{a \times b}{\mathrm{gcd}(a, b)}
$$

Define a function that computes the least common multiple of two integers.
Python (version 3.5 and later) has a `gcd` function you can use.
This cell starts with an import statement that fetches the function from the `math` library.

<!-- grade_id: repeat, points: 6, grade: true, solution: true -->
```python
from math import gcd

def lcm(a,b):
    # YOUR CODE HERE
    lc_multiple = (a*b)/(gcd(a,b))
    return lc_multiple
```

```python
lcm(100,50)
```

Output:
```
100.0
```

```python
lcm(32,8)
```

Output:
```
32.0
```

##### Tests 

The LCM of 4 and 6 is 12:

<!-- grade_id: cell-35159a846e21db3e, points: 2, grade: true, solution: false -->
```python
assert lcm(4,6) == 12
```

The LCM of 21 and 6 is 42:

<!-- grade_id: cell-6dcc9be208bd6f87, points: 2, grade: true, solution: false -->
```python
assert lcm(21,6) == 42
```

###  <span style="color:seagreen">Exercise 2.4: </span>`tax_rate`

A state has defined their income tax rates so that people who earn less than $\$10,000$ pay no taxes, people who earn between $\$10,000$ and $\$19,999$ pay $5\%$, and people who earn $\$20,000$ or more pay $7\%$.

Define a function named `tax_rate` that will take a person's income as an argument and return their tax rate, either 0, 5.0, or 7.0, using the rule described above.

<!-- grade_id: cell-dac2acfcf3298187, points: 7, grade: true, solution: true -->
```python
def tax_rate(x):
    # YOUR CODE HERE
    if x < 10000:
        return round(0,1)
    elif 10000 < x < 19999:
        return 5
    elif x >= 20000:
        return(round(7,1))
```

```python
tax_rate(25000)
```

Output:
```
7
```

```python
tax_rate(1)
```

Output:
```
0
```

##### Tests 

The tax rate for $\$15,000$ income is 5.0.

<!-- grade_id: cell-a198130ab8da8830, points: 1, grade: true, solution: false -->
```python
assert tax_rate(15000) == 5.0
```

The tax rate for $\$5,000$ income is 0.

<!-- grade_id: cell-b0138ad46a65661d, points: 1, grade: true, solution: false -->
```python
assert tax_rate(5000) == 0
```

The tax rate for $\$50,000$ income is 7.0.

<!-- grade_id: cell-a53b1a4573288415, points: 1, grade: true, solution: false -->
```python
assert tax_rate(50000) == 7.0
```

###  <span style="color:seagreen">Exercise 2.5: </span>`pig`

Define a function named `pig` that translates an English word into "Pig Latin".  Some examples:
```
>>> pig('beer')
eerbay

>>> pig('ale')
aleway
```

A complete set of rules can be found on Wikipedia (https://en.wikipedia.org/wiki/Pig_Latin).  For this exercise you only need to implement the two simplest rules:
* If a word starts with a vowel (`a`, `e`, `i`, `o`, or `u`) just append `way` to the end of the word
* Otherwise remove the first letter, move it to the end, and append `ay`


<!-- grade_id: cell-1df4d6f7bee54fba, points: 6, grade: true, solution: true -->
```python
def pig(s):
    # YOUR CODE HERE
    if s[0] in ('a', 'e', 'i', 'o', 'u'):
        return(s +'way')
    else:
        return((s[1:] + s[0]) + 'ay')
```

<!-- grade_id: cell-71c02bfacfca1463, points: 2, grade: true, solution: false -->
```python
assert pig('beer') == 'eerbay'
```

<!-- grade_id: cell-c26c30852b961e0c, points: 2, grade: true, solution: false -->
```python
assert pig('ale') == 'aleway'
```

##### Style Challenge 

Can you write your program so the test for a vowel at the front of a word is done with a single `if` statement? _i.e._ see if you can avoid a “cascade” of `if-elif` statements.

##### Optional Challenge 

Implement some of the other rules described on the Wikipedia page.  If a word starts with more than one consonant all of them should be moved to the end:
```
>>> pig('whiskey')
iskeywhay

>>> pig('string')
ingstray
```